<h1>CS152 Assignment 1: The 8-puzzle</h1>

Before you turn in this assignment, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then run the test cells for each of the questions you have answered.  Note that a grade of 3 for the A* implementation requires all tests in the "Basic Functionality" section to be passed.  The test cells pass if they execute with no errors (i.e. all the assertions are passed).

Make sure you fill in any place that says `YOUR CODE HERE`.  Be sure to remove the `raise NotImplementedError()` statements as you implement your code - these are simply there as a reminder if you forget to add code where it's needed.

---

<h1>
Question 1    
</h1>
Define your <code>PuzzleNode</code> class below.  Ensure that you include all attributes that you need to implement an A* search.  If you wish, you can even include member functions, such as a function to generate successor states.  Alternatively, you can code up this functionality later in the <code>solvePuzzle</code> function.

<!-- grade_id: b5da05, points: 0, grade: false, solution: true -->
```python
# Import any packages you need here
# Also define any variables as needed

# YOUR CODE HERE (OPTIONAL)
import copy
def goal_state(n):
        goal_state = []
        for i in range(n):
            temp = []
            for j in range(n):
                temp.append(n*i + j) 
            goal_state.append(temp)
        return goal_state
        
#Now, define the class PuzzleNode:
class PuzzleNode:
    """
    Class PuzzleNode: Provides a structure for performing A* search for the n^2-1 puzzle
    """
    # YOUR CODE HERE
    def __init__(self,state,n,fval = 0, parent=None,depth=0):
        self.state = state
        self.goal_state = goal_state(n)
        self.n = n
        self.parent = [parent]
        self.children = []
        self.depth = depth
        self.f = fval
    
    def __lt__(self,other):
        return self.f < other.f

    def check_error(self):
        # Check for acceptance of the input problem
        
        # Check if it's an n x n board
        if (len(self.state) != self.n):
            return False
        # Check if each row contain exactly n elements
        if True:
            for i in range(len(self.state)):
                if len(self.state[i]) != self.n:
                    return False
        
        if True:
            temp = []
            for i in range(self.n):
                for j in range(self.n):
                    temp.append(self.state[i][j])
            temp.sort()
            if temp != list(range(0,(self.n)**2)):
                return False

        return True
    
    def create_children(self):
        # Make a move and return a new child configuration
        i,j = self.search_for_blank_space()
        
        i0 = i-1
        j0 = j-1
        i1 = i+1
        j1 = j+1

        if i0 >= 0:
            child1_state = copy.deepcopy(self.state)
            child1_state[i][j] = child1_state[i0][j]
            child1_state[i0][j] = 0
            child1 = PuzzleNode(child1_state,self.n,parent=self,depth=self.depth+1)
            self.children.append(child1)

        if j0 >= 0:
            child2_state = copy.deepcopy(self.state)
            child2_state[i][j] = child2_state[i][j0]
            child2_state[i][j0] = 0
            child2 = PuzzleNode(child2_state,self.n,parent=self,depth=self.depth+1)
            self.children.append(child2)

        if i1 < self.n:
            child3_state = copy.deepcopy(self.state)
            child3_state[i][j] = child3_state[i1][j]
            child3_state[i1][j] = 0
            child3 = PuzzleNode(child3_state,self.n,parent=self,depth=self.depth+1)
            self.children.append(child3)

        if j1 < self.n:
            child4_state = copy.deepcopy(self.state)
            child4_state[i][j] = child4_state[i][j1]
            child4_state[i][j1] = 0
            child4 = PuzzleNode(child4_state,self.n,parent=self,depth=self.depth+1)
            self.children.append(child4)

    def visualize(self):
        # Visualize the state
        for i in range(self.n):
            for j in range(self.n):
                if self.state[i][j] == 0:
                    print("_", end='      ')
                else:
                    if self.state[i][j] > 9:
                        end = "     "
                    else:
                        end = "      "
                    print(self.state[i][j], end=end)
            print("\n")

    def search_for_blank_space(self):
        for i in range(self.n):
            for j in range(self.n):
                if self.state[i][j] == 0:
                    return (i,j)

    def __str__(self):
        return str(self.state)

```

<h1>
Question 2    
</h1>
Define your heuristic functions using the templates below.  Ensure that you extend the <code>heuristics</code> list to include all the heuristic functions you implement.  Note that state will be given as a list of lists, so ensure your function accepts this format.  You may use packages like numpy if you wish within the functions themselves.

<!-- grade_id: 3fb850, points: 0, grade: false, solution: true -->
```python
# Add any additional code here (e.g. for the memoization extension)

# YOUR CODE HERE (OPTIONAL)

# Misplaced tiles heuristic
def h1(state):
    """
    This function returns the number of misplaced tiles, given the board state
    Input:
        -state: the board state as a list of lists
    Output:
        -h: the number of misplaced tiles
    """
    # YOUR CODE HERE
    n = len(state)
    goal = goal_state(n)
    h = 0
    for i in range(n):
        for j in range(n):
            if state[i][j] == 0:
                continue
            elif state[i][j] != goal[i][j]:
                h += 1
    return h

# Manhattan distance heuristic
def h2(state):
    """
    This function returns the Manhattan distance from the solved state, given the board state
    Input:
        -state: the board state as a list of lists
    Output:
        -h: the Manhattan distance from the solved configuration
    """
    # YOUR CODE HERE
    n = len(state)
    h = 0
    for i in range(n):
        for j in range(n):
            if state[i][j] == 0:
                continue
            else:
                # Row difference
                h += abs((state[i][j] // n) - i)
                # Column difference
                h += abs((state[i][j] % n) - j)
    return h
    
# Extra heuristic for the extension.  If implemented, modify the function below
def h3(state):
    """
    This function returns a heuristic that dominates the Manhattan distance, given the board state
    Input:
        -state: the board state as a list of lists
    Output:
        -h: the Heuristic distance of the state from its solved configuration
    """
    return 0

# If you implement more than 3 heuristics, then add any extra heuristic functions onto the end of this list.
heuristics = [h1, h2, h3]
```

<h1>
Question 3    
</h1>
Code up your A* search using the SolvePuzzle function within the template below.  Please do not modify the function header, otherwise the automated testing will fail.  You may define other functions or import packages as needed in this cell or by adding additional cells.

<!-- grade_id: 328d7b, points: 0, grade: false, solution: true -->
```python
# Import any packages or define any helper functions you need here
from queue import PriorityQueue

# YOUR CODE HERE (OPTIONAL)
def check_add(item,open_list,closed_list):
    # Check if we should add the new item into priority queue
    
    # First, check whether a similar state appear in open_list
    
    # check_add((f_value,current_node.children[l].state),frontier.queue,explored_states):
    for value, i, state in open_list:
        # If the configuration has already existed + f_value higher than our item
        # And add the new item to the queue
        # We don't need to remove the larger f_value with same the configuration node
        # Because it'll be place behind the less value one. And before doing the expansion
        # the program will check the explored list and see that we've already checked
        # a similar configuration so it'll ignore that state.
        if state == item[1] and value < item[0]:
            return False

    # Check whether a similar state appear in closed_list
    for value, state in closed_list:
        if state == item[1] and value < item[0]:
            return False

    return True

# Main solvePuzzle function.
def solvePuzzle(state, heuristic):
    """This function should solve the n**2-1 puzzle for any n > 2 (although it may take too long for n > 4)).
    Inputs:
        -state: The initial state of the puzzle as a list of lists
        -heuristic: a handle to a heuristic function.  Will be one of those defined in Question 2.
    Outputs:
        -steps: The number of steps to optimally solve the puzzle (excluding the initial state)
        -exp: The number of nodes expanded to reach the solution
        -max_frontier: The maximum size of the frontier over the whole search
        -opt_path: The optimal path as a list of list of lists.  That is, opt_path[:,:,i] should give a list of lists
                    that represents the state of the board at the ith step of the solution.
        -err: An error code.  If state is not of the appropriate size and dimension, return -1.  For the extention task,
          if the state is not solvable, then return -2
    """
    # YOUR CODE HERE
    # Check if state is one of the appropriate size and dimension
    n = len(state)
    a = PuzzleNode(state,n)
    if a.check_error() == False:    # If the dimension fail to be appropriate
        steps = 0
        exp = 0
        max_frontier = 0
        err = -1
        opt_path = None
        return(steps,exp,max_frontier,opt_path,err)

    # Goal state for this puzzle
    goal = goal_state(n)

    i = 0
    exp = 0
    opt_path = []
    frontier = PriorityQueue()
    frontier.put((0,0,a))
    max_frontier = len(frontier.queue)
    explored = []
    explored_states = []

    # Start the algorithm 

    while not frontier.empty():
        print(frontier)
        current_node = frontier.get()[-1]
        f = current_node.depth + heuristic(current_node.state)
        current_node.create_children()
        
        for l in range(len(current_node.children)):
            
            # Check if the successor is the goal state
            
            if current_node.children[l].state == goal:
                current_node = current_node.children[l]
                steps = current_node.depth
                exp = len(explored)
                err = 0
                opt_path.insert(0,current_node.state)
                # As long as we still have parent node to trace back
                while current_node.parent[0] != None:
                    # Pop that parent node out
                    prev_step = current_node.parent.pop()
                    # Left-append the state into our opt_path list
                    opt_path.insert(0,prev_step.state)
                    # Assign that parent node as our current node
                    current_node = prev_step

                return ([steps,exp,max_frontier,opt_path,err])

            # If not, check if we should add it to the priority queue
            f_value = current_node.children[l].depth + heuristic(current_node.children[l].state)
            if check_add((f_value,current_node.children[l].state),frontier.queue,explored_states):
                # The ID of each state in the queue, to check with one was added to the queue first
                # In case the f_value is the same for two states, we choose expand
                # the node that was in the queue first.
                i += 1
                print(current_node.state, f_value)
                frontier.put((f_value,i,current_node.children[l]))

        
        if len(frontier.queue()) > max_frontier:
            max_frontier = len(frontier.queue)

        explored.append(current_node)
        explored_states.append((f,current_node.state))
        



```

```python
# Import any packages or define any helper functions you need here
from queue import PriorityQueue

# YOUR CODE HERE (OPTIONAL)
def check_add(item,open_list,closed_list):
    # Check if we should add the new item into priority queue

    # First, check whether a similar state appear in open_list
    for value, i, state in open_list:
        # If the configuration has already existed + f_value higher than our item
        # And add the new item to the queue
        # We don't need to remove the larger f_value with same the configuration node
        # Because it'll be place behind the less value one. And before doing the expansion
        # the program will check the explored list and see that we've already checked
        # a similar configuration so it'll ignore that state.
        if state == item[1] and value < item[0]:
            return False

    # Check whether a similar state appear in closed_list
    for value, state in closed_list:
        if state == item[1] and value < item[0]:
            return False

    return True

# Main solvePuzzle function.
def solvePuzzle(state, heuristic):
    """This function should solve the n**2-1 puzzle for any n > 2 (although it may take too long for n > 4)).
    Inputs:
        -state: The initial state of the puzzle as a list of lists
        -heuristic: a handle to a heuristic function.  Will be one of those defined in Question 2.
    Outputs:
        -steps: The number of steps to optimally solve the puzzle (excluding the initial state)
        -exp: The number of nodes expanded to reach the solution
        -max_frontier: The maximum size of the frontier over the whole search
        -opt_path: The optimal path as a list of list of lists.  That is, opt_path[:,:,i] should give a list of lists
                    that represents the state of the board at the ith step of the solution.
        -err: An error code.  If state is not of the appropriate size and dimension, return -1.  For the extention task,
          if the state is not solvable, then return -2
    """
    # YOUR CODE HERE
    # Check if state is one of the appropriate size and dimension
    n = len(state)
    a = PuzzleNode(state,n)
    if a.check_error() == False:    # If the dimension fail to be appropriate
        steps = 0
        exp = 0
        max_frontier = 0
        err = -1
        opt_path = None
        return(steps,exp,max_frontier,opt_path,err)

    # Goal state for this puzzle
    goal = goal_state(n)

    i = 0
    exp = 0
    opt_path = []
    frontier = PriorityQueue()
    frontier.put((0,0,a))
    max_frontier = len(frontier.queue)
    explored = []
    explored_states = []

    # Start the algorithm 

    while not frontier.empty():
        current_node = frontier.get()[-1]
        f = current_node.depth + heuristic(current_node.state)
        current_node.create_children()
        
        for l in current_node.children:
            
            # Check if the successor is the goal state
            
            if l.state == goal:
                current_node = l
                steps = current_node.depth
                exp = len(explored)
                err = 0
                opt_path.insert(0,current_node.state)
                # As long as we still have parent node to trace back
                while current_node.parent[0] != None:
                    # Pop that parent node out
                    prev_step = current_node.parent.pop()
                    # Left-append the state into our opt_path list
                    opt_path.insert(0,prev_step.state)
                    # Assign that parent node as our current node
                    current_node = prev_step

                return ([steps,exp,max_frontier,opt_path,err])

            # If not, check if we should add it to the priority queue
            f_value = l.depth + heuristic(l.state)
            if check_add((f_value,l.state),frontier.queue,explored_states):
                # The ID of each state in the queue, to check with one was added to the queue first
                # In case the f_value is the same for two states, we choose expand
                # the node that was in the queue first.
                i += 1
                frontier.put((f_value,i,l))

        if len(frontier.queue) > max_frontier:
            max_frontier = len(frontier.queue)

        explored.append(current_node)
        explored_states.append((f,current_node.state))
        



```

<h1>Extension Questions</h1>

The extensions can be implemented by modifying the code from Q2-3 above appropriately.

1. <b>Initial state solvability:</b>  Modify your SolvePuzzle function code in Q3 to return -2 if an initial state is not solvable to the goal state.
2. <b>Extra heuristic function:</b> Add another heuristic function (e.g. pattern database) that dominates the misplaced tiles and Manhattan distance heuristics to your Q2 code.
3. <b>Memoization:</b>  Modify your heuristic function definitions in Q2 by using a Python decorator to speed up heuristic function evaluation

There are test cells provided for extension questions 1 and 2.

<h1>Basic Functionality Tests</h1>
The cells below contain tests to verify that your code is working properly to be classified as basically functional.  Please note that a grade of <b>3</b> on #aicoding and #search as applicable for each test requires the test to be successfully passed.  <b>If you want to demonstrate some other aspect of your code, then feel free to add additional cells with test code and document what they do.<b>

<!-- grade_id: 8fbc24, points: 1, grade: true, solution: false -->
```python
## Test for state not correctly defined

incorrect_state = [[0,1,2],[2,3,4],[5,6,7]]
_,_,_,_,err = solvePuzzle(incorrect_state, lambda state: 0)
assert(err == -1)
```

<!-- grade_id: cell-5738a2124877e52e, points: 2, grade: true, solution: false -->
```python
## Heuristic function tests for misplaced tiles and manhattan distance

# Define the working initial states
working_initial_states_8_puzzle = ([[2,3,7],[1,8,0],[6,5,4]], [[7,0,8],[4,6,1],[5,3,2]], [[5,7,6],[2,4,3],[8,1,0]])

# Test the values returned by the heuristic functions
h_mt_vals = [7,8,7]
h_man_vals = [15,17,18]

for i in range(0,3):
    h_mt = heuristics[0](working_initial_states_8_puzzle[i])
    h_man = heuristics[1](working_initial_states_8_puzzle[i])
    assert(h_mt == h_mt_vals[i])
    assert(h_man == h_man_vals[i])

```

<!-- grade_id: cell-99c07d9541559b52, points: 4, grade: true, solution: false -->
```python
## A* Tests for 3 x 3 boards
## This test runs A* with both heuristics and ensures that the same optimal number of steps are found
## with each heuristic.

# Optimal path to the solution for the first 3 x 3 state
opt_path_soln = [[[2, 3, 7], [1, 8, 0], [6, 5, 4]], [[2, 3, 7], [1, 8, 4], [6, 5, 0]], 
                 [[2, 3, 7], [1, 8, 4], [6, 0, 5]], [[2, 3, 7], [1, 0, 4], [6, 8, 5]], 
                 [[2, 0, 7], [1, 3, 4], [6, 8, 5]], [[0, 2, 7], [1, 3, 4], [6, 8, 5]], 
                 [[1, 2, 7], [0, 3, 4], [6, 8, 5]], [[1, 2, 7], [3, 0, 4], [6, 8, 5]], 
                 [[1, 2, 7], [3, 4, 0], [6, 8, 5]], [[1, 2, 0], [3, 4, 7], [6, 8, 5]], 
                 [[1, 0, 2], [3, 4, 7], [6, 8, 5]], [[1, 4, 2], [3, 0, 7], [6, 8, 5]], 
                 [[1, 4, 2], [3, 7, 0], [6, 8, 5]], [[1, 4, 2], [3, 7, 5], [6, 8, 0]], 
                 [[1, 4, 2], [3, 7, 5], [6, 0, 8]], [[1, 4, 2], [3, 0, 5], [6, 7, 8]], 
                 [[1, 0, 2], [3, 4, 5], [6, 7, 8]], [[0, 1, 2], [3, 4, 5], [6, 7, 8]]]

astar_steps = [17, 25, 28]
for i in range(0,3):
    steps_mt, expansions_mt, _, opt_path_mt, _ = solvePuzzle(working_initial_states_8_puzzle[i], heuristics[0])
    steps_man, expansions_man, _, opt_path_man, _ = solvePuzzle(working_initial_states_8_puzzle[i], heuristics[1])
    # Test whether the number of optimal steps is correct and the same
    assert(steps_mt == steps_man == astar_steps[i])
    # Test whether or not the manhattan distance dominates the misplaced tiles heuristic in every case
    assert(expansions_man < expansions_mt)
    # For the first state, test that the optimal path is the same
    if i == 0:
        assert(opt_path_mt == opt_path_soln)

```

Output:
```
<ipython-input-9-70b40d63f2ef> in solvePuzzle(state, heuristic)
     93             # If not, check if we should add it to the priority queue
     94             f_value = current_node.children[l].depth + heuristic(current_node.children[l].state)
---> 95             if check_add((f_value,current_node.children[l].state),frontier.queue,explored_states):
     96                 # The ID of each state in the queue, to check with one was added to the queue first
     97                 # In case the f_value is the same for two states, we choose expand

<ipython-input-9-70b40d63f2ef> in check_add(item, open_list, closed_list)
     18 
     19     # Check whether a similar state appear in closed_list
---> 20     for value, state in closed_list:
     21         if state == item[1] and value < item[0]:
     22             return False

KeyboardInterrupt:
```

<!-- grade_id: cell-2981ba2cc7745c22, points: 8, grade: true, solution: false -->
```python
## A* Test for 4 x 4 board
## This test runs A* with both heuristics and ensures that the same optimal number of steps are found
## with each heuristic.

working_initial_state_15_puzzle = [[1,2,6,3],[0,9,5,7],[4,13,10,11],[8,12,14,15]]
steps_mt, expansions_mt, _, _, _ = solvePuzzle(working_initial_state_15_puzzle, heuristics[0])
steps_man, expansions_man, _, _, _ = solvePuzzle(working_initial_state_15_puzzle, heuristics[1])
# Test whether the number of optimal steps is correct and the same
assert(steps_mt == steps_man == 9)
# Test whether or not the manhattan distance dominates the misplaced tiles heuristic in every case
assert(expansions_mt >= expansions_man)

```

```python

```

Output:
```
8
```

<h1>Extension Tests</h1>
The cells below can be used to test the extension questions.  Memoization if implemented will be tested on the final submission - you can test it yourself by testing the execution time of the heuristic functions with and without it.

<!-- grade_id: cell-21cada9978e9a1bb, points: 16, grade: true, solution: false -->
```python
## Puzzle solvability test

unsolvable_initial_state = [[7,5,6],[2,4,3],[8,1,0]]
_,_,_,_,err = solvePuzzle(unsolvable_initial_state, lambda state: 0)
assert(err == -2)
```

<!-- grade_id: cell-8c7a414efc152379, points: 32, grade: true, solution: false -->
```python
## Extra heuristic function test.  
## This tests that for all initial conditions, the new heuristic dominates over the manhattan distance.

dom = 0
for i in range(0,3):
    steps_new, expansions_new, _, _, _ = solvePuzzle(working_initial_states_8_puzzle[i], heuristics[2])
    steps_man, expansions_man, _, _, _ = solvePuzzle(working_initial_states_8_puzzle[i], heuristics[1])
    # Test whether the number of optimal steps is correct and the same
    assert(steps_new == steps_man == astar_steps[i])
    # Test whether or not the manhattan distance is dominated by the new heuristic in every case, by checking
    # the number of nodes expanded
    dom = expansions_man - expansions_new
assert(dom > 0)
```

<!-- grade_id: cell-09f710b6aa2e3fa3, points: 64, grade: true, solution: false -->
```python
## Memoization test - will be carried out after submission
```